# Decision Trees

### Concepts
You're given a dataset of **30** elements, 15 of which belong to a positive class (denoted by `+` ) and 15 of which do not (denoted by `-`). These elements are described by two attributes, A and B, that can each have either one of two values, true or false.

The diagrams below show the result of splitting the dataset by attribute: the diagram on the left hand side shows that if we split by attribute A there are 13 items of the positive class and 2 of the negative class in one branch and 2 of the positive and 13 of the negative in the other branch. The right hand side shows that if we split the data by attribute B there are 8 items of the positive class and 7 of the negative class in one branch and 7 of the positive and 8 of the negative in the other branch.

<img src="./decision_stump.png">

### 1) Which one of the two attributes resulted in the best split of the original data?

Assign `ans1` to either `"A"` or `"B"` to represent which attribute is best.

<!-- grade_id: cell-7f0773ccaaca4a5e, points: 0, grade: false, solution: true -->
```python
# Replace None with appropriate code
# entropy = -p_positive * log2(p_positive) - p_negative * log2(p_negative)
# entropy_before_split = -15/30*log2(15/30)-15/30*log2(15/30) = 1


 # Determine the best attribute
# ans1 = "A" if info_gain_A > info_gain_B else "B"
# YOUR CODE HERE
#raise NotImplementedError()
# import math
import math

# Entropy calculation function
def entropy(p_positive, p_negative):
    return -p_positive * math.log2(p_positive) - p_negative * math.log2(p_negative)

# Entropy before split
entropy_before_split = entropy(0.5, 0.5)

# For Attribute A
p_positive_A = 13 / 15
p_negative_A = 2 / 15
entropy_after_split_A = entropy(p_positive_A, p_negative_A)

# For Attribute B
p_positive_B = 8 / 15
p_negative_B = 7 / 15
entropy_after_split_B = entropy(p_positive_B, p_negative_B)

# Weighted average entropy
weighted_avg_entropy_A = (15 / 30) * entropy_after_split_A + (15 / 30) * entropy_after_split_A
weighted_avg_entropy_B = (15 / 30) * entropy_after_split_B + (15 / 30) * entropy_after_split_B

# Information gain
info_gain_A = entropy_before_split - weighted_avg_entropy_A
info_gain_B = entropy_before_split - weighted_avg_entropy_B

# Determine the best attribute
ans1 = "A" if info_gain_A > info_gain_B else "B"

ans1



```

Output:
```
'A'
```

<!-- grade_id: cell-e31f81fb8b1e2fbc, points: 1, grade: true, solution: false -->
```python
# PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
# THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS

# ans1 should be the string "A" or "B"
assert type(ans1) == str

```

### Decision Trees for Regression

In this section, you will use decision trees to fit a regression model to the Combined Cycle Power Plant dataset.

This dataset is from the UCI ML Dataset Repository, and has been included in the `data` folder of this repository as an Excel `.xlsx` file, `Folds5x2_pp.xlsx`.

The features of this dataset consist of hourly average ambient variables taken from various sensors located around a power plant that record the ambient variables every second.  
- Temperature (AT)
- Ambient Pressure (AP)
- Relative Humidity (RH)
- Exhaust Vacuum (V)

The target to predict is the net hourly electrical energy output (PE).

The features and target variables are not normalized.

In the cells below, we import `pandas` and `numpy` for you, and we load the data into a pandas DataFrame. We also include code to inspect the first five rows and get the shape of the DataFrame.

```python
# Run this cell without changes

import pandas as pd
import numpy as np

# Load the data
filename = 'Folds5x2_pp.csv'
df = pd.read_csv(filename, index_col=0)
```

```python
# Run this cell without changes
df.head()
```

Output:
```
AT      V       AP     RH      PE
0  14.96  41.76  1024.07  73.17  463.26
1  25.18  62.96  1020.04  59.08  444.37
2   5.11  39.40  1012.16  92.14  488.56
3  20.86  57.32  1010.24  76.64  446.48
4  10.82  37.50  1009.23  96.62  473.90
```

```python
# Run this cell without changes
df.shape
```

Output:
```
(9568, 5)
```

Below, we split the data into features and target ('PE') for you.

```python
# Run this cell without changes
X = df[df.columns.difference(['PE'])]
y = df['PE']
```

### 2) Use `train_test_split` ([documentation here](https://scikit-learn.org/stable/modules/generated/sklearn.model_selection.train_test_split.html)) to split the data into training and test sets. Create training and test sets with `test_size=0.5` and `random_state=1`.

You will need to import the relevant code from scikit-learn

<!-- grade_id: 0d38786c25982cb53f9a4a1b6e96c717, points: 0, grade: false, solution: true -->
```python
# Replace None with appropriate code

# Import the relevant function
from sklearn.model_selection import train_test_split


# Split the data
X_train, X_test, y_train, y_test =train_test_split(X, y, test_size=0.5, random_state=1)


# YOUR CODE HERE
#raise NotImplementedError()


```

<!-- grade_id: d24a4d515abed049b603cf57d0621c24, points: 1, grade: true, solution: false -->
```python
# PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
# THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS

# X_train and X_test should be dataframes
assert type(X_train) == pd.DataFrame and type(X_test) == pd.DataFrame

# y_train and y_test should be series
assert type(y_train) == pd.Series and type(y_test) == pd.Series

```

### 3) Fit a decision tree regression model with scikit-learn to the training data. Use parameter defaults and `random_state=1` for this model

For the rest of this section feel free to refer to the scikit-learn documentation on [decision tree regressors](https://scikit-learn.org/stable/modules/generated/sklearn.tree.DecisionTreeRegressor.html).

You will need to import the relevant code from scikit-learn

Assign your regressor (model) to the variable `dt`

<!-- grade_id: a9209fbe3444034fdaa3f956a506fbff, points: 0, grade: false, solution: true -->
```python
# Replace None with appropriate code

# Import the relevant class
from sklearn.tree import DecisionTreeRegressor

# Instantiate the model
dt = DecisionTreeRegressor(random_state=1)

# Fit the model to the training data here
dt.fit(X_train, y_train)

# YOUR CODE HERE
#raise NotImplementedError()

# Testing out the model's r2 score on the training data overall
dt_train_score = dt.score(X_train, y_train)
dt_train_score
```

Output:
```
1.0
```

<!-- grade_id: 86c405061165b88b5cfc0b624e3cb7ab, points: 0.5, grade: true, solution: false -->
```python
# PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
# THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS

# dt should be a decision tree regressor
assert type(dt) == DecisionTreeRegressor

```

### 4) That score looks good, but are we overfitting? Obtain the cross-validated coefficient of determination (r2 score) of the predictions of `dt` on the training data

You can use the default `cv` value of 5, and you do not need to specify a custom scorer, since r2 is the default for a `DecisionTreeRegressor`

You will need to import the relevant code from scikit-learn

Assign your answer to the variable `dt_cv`

<!-- grade_id: f029a5f5197e376c5c4d067faf83d300, points: 0, grade: false, solution: true -->
```python
# Replace None with appropriate code

# Import the relevant function
from sklearn.model_selection import cross_val_score

# Assign the cross validated score to dt_cv
dt_cv = cross_val_score(dt, X_train, y_train, cv=5, scoring='r2')

# YOUR CODE HERE
#raise NotImplementedError()

dt_cv
```

Output:
```
array([0.90892482, 0.90764959, 0.92975025, 0.90521069, 0.91216957])
```

<!-- grade_id: b6f7124b3dbeec9dc9dedf01a4cef54e, points: 0.5, grade: true, solution: false -->
```python
# PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
# THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS

# dt_cv should be a NumPy array
assert type(dt_cv) == np.ndarray
# there should be 5 scores in dt_cv
assert len(dt_cv) == 5

```

### Hyperparameter Tuning of Decision Trees for Regression

### 5) Create a second decision tree model, this time with the following hyperparameters specified:
 - `random_state=1`
 - `max_depth=5`

### Fit it to the training data

Assign the new model to the variable `dt_tuned`

<!-- grade_id: 92bf6328b83bcf32a2852fb7800e594f, points: 0, grade: false, solution: true -->
```python
# Replace None with appropriate code

# Create a second decision tree model
dt_tuned = DecisionTreeRegressor(random_state=1, max_depth=5)

# Fit the new model on the training data
dt_tuned.fit(X_train, y_train)


# YOUR CODE HERE
#raise NotImplementedError()

# Testing out the model's r2 score on the training data overall
dt_tuned_train_score = dt_tuned.score(X_train, y_train)
dt_tuned_train_score
```

Output:
```
0.9365685127196766
```

<!-- grade_id: 29b36b6337380071c558d458823a426a, points: 0.5, grade: true, solution: false -->
```python
# PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
# THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS

# dt_tuned sould be a DecisionTreeRegressor
assert type(dt_tuned) == DecisionTreeRegressor

```

### 6) Now the score on the training data overall is worse, but how does this generalize to unseen data? Obtain the cross-validated coefficient of determination (r2 score) of the predictions of `dt_tuned` on the training data

Use the same arguments as the previous cross validation

Assign your answer to the variable `dt_tuned_cv`

<!-- grade_id: ffea576577891e8a5cb4df641a1697d4, points: 0, grade: false, solution: true -->
```python
# Replace None with appropriate code
dt_tuned_cv = cross_val_score(dt_tuned, X_train, y_train, cv=5, scoring='r2')

# YOUR CODE HERE
#raise NotImplementedError()

dt_tuned_cv
```

Output:
```
array([0.93172053, 0.93093166, 0.93536142, 0.91836297, 0.92773155])
```

<!-- grade_id: 028ded720a13491278d805d468432521, points: 0.5, grade: true, solution: false -->
```python
# PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
# THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS

# dt_tuned_cv should be a NumPy array
assert type(dt_tuned_cv) == np.ndarray
# there should be 5 scores in dt_tuned_cv
assert len(dt_tuned_cv) == 5

```

```python
# Run this cell without changes:
print("Train score for dt:      ", dt_train_score)
print("Train score for dt_tuned:", dt_tuned_train_score)
print()
print("CV scores for dt:      ", dt_cv)
print("CV scores for dt_tuned:", dt_tuned_cv)
```

Output:
```
Train score for dt:       1.0
Train score for dt_tuned: 0.9365685127196766

CV scores for dt:       [0.90892482 0.90764959 0.92975025 0.90521069 0.91216957]
CV scores for dt_tuned: [0.93172053 0.93093166 0.93536142 0.91836297 0.92773155]
```

### 7) Assuming you want a generalizable model that will perform well on future, unseen data, which model has better performance?  

Assign `ans7` to either `"dt"` or `"dt_tuned"` to represent which model is best.

<!-- grade_id: cell-feffc934dc6415f8, points: 0, grade: false, solution: true -->
```python
ans7 = "dt_tuned"
# YOUR CODE HERE
#raise NotImplementedError()
ans7
```

Output:
```
'dt_tuned'
```

<!-- grade_id: 5305c07427b2a67d7dcb6cb40f83785c, points: 1, grade: true, solution: false -->
```python
assert type(ans7) == str
assert ans7 in ['dt', 'dt_tuned']
```

```python
#A model with consistently high cross-validated scores is likely to generalize better.
#The tuned decision tree (dt_tuned) has consistently high cross-validated scores across folds,
#indicating better generalization performance compared to the default decision tree (dt).
#Therefore, we can assign "dt_tuned" as the model with better performance.
```